# Denison CS181/DA210

---

## Filesystem Exercises

Execute the prolog cell, and then complete the following exercises reviewing working with lists.

Note that **all functions that you submit must have docstrings**.  Points will be deducted for missing docstrings.

```python
import os
import io
import sys
from contextlib import redirect_stdout
from IPython.core.debugger import set_trace

def add_modules():
    """
    Starting at the current directory and proceeding up the file system
    tree, search for a directory named `modules`.  If found, and if not
    already there, add to the Python module search path.
    
    Params: None
    
    Return: None
    """
    directory = "."
    levels = 0
    while not os.path.isdir(os.path.join(directory, "modules")) and \
          levels < 5:
        directory = os.path.join(directory, "..")
        levels += 1
    module_path = os.path.abspath(os.path.join(directory, "modules"))
    if os.path.isdir(module_path):
        if not module_path in sys.path:
            sys.path.append(module_path)

add_modules()
import util

datadir = util.resolve_dir("otherdata")
```

**Q1** Write a function

    parentDir()
    
that obtains the absolute path for the parent directory/folder relative to the execution directory of the current notebook.  While it is permissible to change the current working directory as part of your solution, upon completion of the function, the current working directory **must be** the same as it was before the function was called.

<!-- grade_id: cell-af780bee653c4cef, points: 0, grade: false, solution: true -->
```python
def parentDir():
    """
    This function gets the directory that your current directory is in.
    It will end in the directory the user started in after running.
    
    Params: None
    
    Return: The directory that your current directory is in.
    """
    startDir = os.getcwd()
    os.chdir("..")
    newDir = os.getcwd()
    os.chdir(startDir)
    return newDir
    
```

<!-- grade_id: cell-58df57fbed4ae45f, points: 2, grade: true, solution: false -->
```python
pdir = parentDir()
assert os.path.isdir(pdir)
assert os.path.isdir(os.path.join(parentDir(), "HW_1.1"))
```

**Q2** Just like we often search for an element in a collection, we often might want to search and determine whether a particular file or directory is present at a particular location in the file system.

Write a function

    searchForFolder(dir, keyfolder)
    
that searches for a folder, `keyfolder` in the direcory given by `dir`.  If the folder is found, the string **path** for the folder is returned.  This path can be either relative or absolute.  If the name in `keyfolder` is not present, or if it is not a **directory**, the function should return `None` to so indicate.  (Note that `None` is the Python "empty" type; it is not a string with those characters.)

<!-- grade_id: cell-5a2690ad91df7a9f, points: 0, grade: false, solution: true -->
```python
def searchForFolder(dir, keyfolder):
    """
    This function searches a given directory for a given folder
    name.
    
    Params: dir: the directory that is going to be searched.
            keyfolder: the folder that is being searched for.
    
    Return: searchPath: will return the path of the folder if
                        the folder is found
            None: will return none if the folder is not found.
    """
    searchPath = os.path.join(dir, keyfolder)
    if os.path.isdir(searchPath):
        return searchPath
    else:
        return None
```

<!-- grade_id: cell-07f7d7e0c1519cdf, points: 2, grade: true, solution: false -->
```python
# Testing cell
test1 = searchForFolder(".", "testfolder")
assert isinstance(test1, str)
assert os.path.isdir(test1)

test2 = searchForFolder(os.path.join(os.getcwd(), ".."), 
                        "HW_1.1")
assert isinstance(test2, str)
assert os.path.isdir(test2)

test3 = searchForFolder(".", "lists.ipynb")
assert test3 == None

test4 = searchForFolder(".", "foobar")
assert test4 == None
```

**Q** Write a function

    countFileMatch(dir, prefix)
    
that searches the files in the give directory, `dir`, and determines the **count** of the number of files where the start of the file name matches `prefix`.  It should not count any folders whose name might match.  If no files in the directory match the prefix, the function should return 0.  If the directory does not exist or the variable refers to something that is not a directory, the function should return `None`.

<!-- grade_id: cell-c545311236841784, points: 0, grade: false, solution: true -->
```python
def countFileMatch(dir, prefix):
    """
    This function counts the amount of files in a given directory that
    start with a given prefix. The amount of prefix files gets returned.
    
    Params: dir: the directory that is being searched.
            prefix: the prefix for files that is getting searched for.
    
    Return: numPre: the amount of files that start with the prefix
            None: will return None if the directory given does not
                  exist.
    """
    if os.path.isdir(dir):
        itemList = os.listdir(dir)
        numPre = 0;
        for i in range(len(itemList)):
            if prefix == itemList[i][0:len(prefix)]:
                if os.path.isfile(os.path.join(dir, itemList[i])):
                    numPre = numPre + 1
        return numPre
    else:
        return None
    
```

<!-- grade_id: cell-bdb26bc354eb4f3b, points: 3, grade: true, solution: false -->
```python
assert countFileMatch(".", "file") == 2
assert countFileMatch("..", "HW") == 0
assert countFileMatch(".", "file") == 2
assert countFileMatch(os.getcwd(), "test") == 0
assert countFileMatch("lists.ipynb", "list") == None
assert countFileMatch(os.path.join(os.getcwd(), "..", "foobar"), "") == None
assert countFileMatch(".", "") == 3
```

```python

```