# Denison CS181/DA210

---

## List Review Exercises

Execute the prolog cell, and then complete the following exercises reviewing working with lists.

Note that **all functions that you submit must have docstrings**.  Points will be deducted for missing docstrings.

```python
import os
import io
import sys
from contextlib import redirect_stdout
from IPython.core.debugger import set_trace

def add_modules():
    """
    Starting at the current directory and proceeding up the file system
    tree, search for a directory named `modules`.  If found, and if not
    already there, add to the Python module search path.
    
    Params: None
    
    Return: None
    """
    directory = "."
    levels = 0
    while not os.path.isdir(os.path.join(directory, "modules")) and \
          levels < 5:
        directory = os.path.join(directory, "..")
        levels += 1
    module_path = os.path.abspath(os.path.join(directory, "modules"))
    if os.path.isdir(module_path):
        if not module_path in sys.path:
            sys.path.append(module_path)

add_modules()
import util

datadir = util.resolve_dir("otherdata")
```

**Q1** Incorporating a `range()` function expression, create a list of the even integers less than or equal to 14 and greater than 0, in descending order, assigning to `L1`.  Make sure the list is fully generated to integer elements, and thus `L1` is a real list and not a reference to a `range` object.

<!-- grade_id: cell-4e65eeb626916cf0, points: 0, grade: false, solution: true -->
```python
# Solution cell

L1 = []
i = 14
while i > 0:
    L1.append(i)
    i = i - 2
```

Output:
```
[14, 12, 10, 8, 6, 4, 2]
```

<!-- grade_id: cell-241c20fdf54b9430, points: 1, grade: true, solution: false -->
```python
assert len(L1) == 7
for element in L1:
    assert isinstance(element, int)
assert L1[0] == 14
assert L1[len(L1)-1] == 2
```

**Q2** Construct a seven element list referred to by `L2` whose elements are all real-valued numbers of your choosing.  Then construct list `L3` whose values are, by index position, the value of each element from `L2` divided by the corresponding element from `L1`.

<!-- grade_id: cell-50159efdc835e09c, points: 0, grade: false, solution: true -->
```python
# Solution cell

L2 = [2.0,5.5,6.4,4.2,3.6,2.2,1.0]
L3 = []
k=0
for k in range(7):
        L3.append(L2[k]/L1[k])
print(L2)
print(L3)
```

Output:
```
[2.0, 5.5, 6.4, 4.2, 3.6, 2.2, 1.0]
[0.14285714285714285, 0.4583333333333333, 0.64, 0.525, 0.6, 0.55, 0.5]
```

<!-- grade_id: cell-ecd6277a7553b198, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert len(L2) == 7
assert len(L3) == 7
for element in L2:
    assert isinstance(element, float)
    
# In the instructor's version, there follow hidden tests
# to validate that the values in L3 match the specification

```

**Execute the following cell prior to answering the questions in this section.**

```python
import random 
random.seed(19610218)

L1 = [random.randint(-20,20) for sample in range(15)]
L2 = [random.randint(-20,20) for sample in range(15)]
```

**Q3** Perform the following assignments accessing the specified elements of `L1` and `L2`:

- Assign to `v1` the value of the third element in list `L1`.
- Assign to `v2` the value of the second to last element of list `L2`.
- Assign to `S1` the sublist from the beginning of `L2` up to index 5 (inclusive).
- Assign to `S2` the sublist from the fourth from the end of `L1` up to (but not including) the last element in that list.
- Assign to `S3` the sublist from the fifth element up to the end of `L2`, selecting every third element.

<!-- grade_id: cell-9c7b3ca62d1969f6, points: 0, grade: false, solution: true -->
```python
print('L1:', L1)
print('L2:', L2)
v1 = L1[2]
v2 = L2[13]
S1 = L2[0:6]
S2 = L1[11:14]
S3 = []
for i in range(4,16,3):
    S3.append(L2[i])
print('v1:', v1, 'v2:', v2)
print('S1:', S1)
print('S2:', S2)
print('S3:', S3)
```

Output:
```
L1: [-1, 4, -17, 8, -2, 15, -18, -5, -4, 16, 3, -9, -17, 3, -14]
L2: [12, 20, -8, 2, -10, 1, 14, 2, 5, 14, 11, 18, 10, 2, 13]
v1: -17 v2: 2
S1: [12, 20, -8, 2, -10, 1]
S2: [-9, -17, 3]
S3: [-10, 2, 11, 2]
```

<!-- grade_id: cell-1e7b76a3fbb1f9cd, points: 3, grade: true, solution: false -->
```python
assert v1 == -17
assert v2 == 2
assert S1 == [12, 20, -8, 2, -10, 1]
assert S2 == [-9, -17, 3]
assert S3 == [-10, 2, 11, 2]
```

**Q4** Write a function

    listDecrement(plist)
    
that modifies the list given by `plist` so that each value is one less than before.  This function should work **in place** and mutate the given list.

<!-- grade_id: cell-7908ddf6259e6de5, points: 0, grade: false, solution: true -->
```python
def listDecrement(plist):
    """
    This function takes a list and decreases each of the elements
    by 1.
    
    Params: plist: the list you want to decrease each element by 1
    
    Return: None
    """
    for i in range(len(plist)):
        plist[i] = plist[i] - 1

```

<!-- grade_id: cell-6170c362782138b3, points: 2, grade: true, solution: false -->
```python
L1 = []
listDecrement(L1)
assert L1 == []

funout = io.StringIO()
with redirect_stdout(funout):
    listDecrement([5,4,3,2,1])
assert funout.getvalue() == ""

L2 = [5,4,3,2,1]
listDecrement(L2)
assert L2 == [4,3,2,1,0]

assert listDecrement([5]) == None
```

**Q5** Write a function

    listDecrease(plist, amount)
    
that creates and returns a new list where values are determined by subtracting `amount` from each numeric value in `plist`.  **If an element in `plist` is neither an integer (int) nor a real-valued number (float), the element should be skipped over.**  So this function is an example of both accumulation and filtering list patterns.

You may want to look up the function `isinstance` to find out how to determine if a variable is the same (or a subclass) of a specific type.  The `isinstance()` function is often used in homework testing cells.

<!-- grade_id: cell-ca43062d63ddb6e1, points: 0, grade: false, solution: true -->
```python
def listDecrease(plist, amount):
    """
    This function creates a new list using another list and decreasing
    its values by a given amount. It will do the integers and the floats 
    of the list but will skip over any other data type.
    
    Params: plist: the list you want to decrease the values of.
            amount: the amount you want to decrease the lists values of.
    
    Return: newList: the new decreased list
    """
    newList = []
    for i in range(len(plist)):
        if isinstance(plist[i], float) or isinstance(plist[i], int):
            newList.append(plist[i] - amount)
    return newList
```

<!-- grade_id: cell-941603e9308da167, points: 2, grade: true, solution: false -->
```python
assert listDecrease([], 0) == []
assert listDecrease([1], 1) == [0]
assert listDecrease([5, 1, 3, 4, 2], 5) == [0, -4, -2, -1, -3]
```

**Q6** Implement a function

    mySearch(plist, key)
    
that searches `plist` for element instances that match `key`.  The function should **build a tuple** whose value contains the **set of indices** for where matches are found.  For instance, `mySearch([3, 5, 1, 5, 2], 5)` should return the tuple `(1, 3)`, while `mySearch([3, 5, 1, 5, 2], 4)` should return an empty tuple `()`.

<!-- grade_id: cell-8e3ea24bd4b08307, points: 0, grade: false, solution: true -->
```python
def mySearch(plist, key):
    """
    This function searches an entire list for intergers that
    match a key integer. Once if finds one it will add its index
    to a tuple and continue searching. It will return tuple at
    the end of its search.
    
    Params: plist: the list that is getting searched.
            key: the integer that is being looked for.
    
    Return: indTup: the tuple with the indices of they keys in
                    the list.
    """
    indTup = () 
    for i in range(len(plist)):
        if plist[i] == key:
            indTup = indTup + (i,)
    return indTup
```

<!-- grade_id: cell-609c87bc561dff9d, points: 2, grade: true, solution: false -->
```python
assert mySearch([], 0) == ()
assert mySearch([3, 5, 1, 5, 2], 5) == (1, 3)
assert mySearch([3, 5, 1, 5, 2], 4) == ()
assert mySearch([3, 5, 1, 5, 2], 3) == (0,)
assert mySearch([3, 5, 1, 5, 2], 2) == (4,)
```

```python

```