# Denison CS-181/DA-210 Homework

---

## Functions and Lambda Expressions

*Execute the prolog cell*

```python
import os
import io
import sys
from contextlib import redirect_stdout

def add_modules():
    """
    Starting at the current directory and proceeding up the file system
    tree, search for a directory named `modules`.  If found, and if not
    already there, add to the Python module search path.
    
    Params: None
    
    Return: None
    """
    directory = "."
    levels = 0
    while not os.path.isdir(os.path.join(directory, "modules")) and \
          levels < 5:
        directory = os.path.join(directory, "..")
        levels += 1
    module_path = os.path.abspath(os.path.join(directory, "modules"))
    if os.path.isdir(module_path):
        if not module_path in sys.path:
            sys.path.append(module_path)

add_modules()
import util
```

**Q1** Write a lambda function that takes three parameters and multiplies them together.  Assign the lambda function expression to `mult3`.

<!-- grade_id: cell-c3740a7586d5d621, points: 0, grade: false, solution: true -->
```python
# Solution cell

mult3 = lambda x, y, z : x*y*z
mult3(10,12, 9)
```

Output:
```
1080
```

<!-- grade_id: cell-26c960f88ccba2cb, points: 1, grade: true, solution: false -->
```python
assert mult3(10, 12, 9) == 1080
```

**Q2** Write a lambda expression and assign it to the variable `oddcubed` such that it cubes the value of its single argument if the argument is odd, and doubles it if the argument if it is even.

Afterward, for example, expression `oddcubed(5)` should yield `125`, while `oddcubed(4)` should yield `8`.

Hint: We need an *expression* to handle an `if` type condition.  See the book discussion on the ternary `if`.

<!-- grade_id: cell-2d5fb1f763401dd7, points: 0, grade: false, solution: true -->
```python
# Solution cell

oddcubed = lambda x : x**3 if x % 2 != 0 else 2*x

print(oddcubed(5))
print(oddcubed(4))

```

Output:
```
125
8
```

<!-- grade_id: cell-a02c95146482fa97, points: 2, grade: true, solution: false -->
```python
assert oddcubed(5) == 125
assert oddcubed(0) == 0
assert oddcubed(-1) == -1
assert oddcubed(2) == 4
assert oddcubed(-2) == -4
```

**Q3** Write a function

`whichIsSmaller(f1,f2,n)`

which, in the body of the function, computes `f1(n)` and `f2(n)`, and returns a reference to the function (`f1` or `f2`) which results in the smaller resultant value when an integer `n` is passed as a parameter. In case of a tie, return the first passed function (`f1`). Test with `fun1(x)` defined as $x^3$ and `fun2(x)` defined as $5x^2+4$ and using $n=2$. Test again with $n=100$.

<!-- grade_id: cell-456f1456d974fe9f, points: 0, grade: false, solution: true -->
```python
# Solution cell
def whichIsSmaller(f1,f2,n):
    '''
    This function takes in two functions and an integer, puts
    the integer in the two functions and which ever results
    in a smaller number will be referenced.
    
    Params: f1: The first function you will put an integer in.
            f2: The second function you will put an integer in.
            n: The integer you will pass through the two functions.
    
    Returns: A reference to the function that outputs the smaller
             result.
    '''
    if f1(n) <= f2(n):
        return f1
    return f2

def fun1(x):
    return x**3

def fun2(x):
    return 5*x*x + 4

print(whichIsSmaller(fun1,fun2,2))
print(whichIsSmaller(fun1,fun2,100))
```

Output:
```
<function fun1 at 0x000002C389BDB9D0>
<function fun2 at 0x000002C389BDB940>
```

<!-- grade_id: cell-dd38486315718564, points: 3, grade: true, solution: false -->
```python
# Testing cell

def fun1(x):
    return x**3 - 7

def fun2(x):
    return 5*x*x + 4*x + 2

assert whichIsSmaller(fun1,fun2,0) == fun1
assert whichIsSmaller(fun1,fun2,-10) == fun1
assert whichIsSmaller(fun1,fun2,2) == fun1
assert whichIsSmaller(fun1,fun2,50) == fun2

```

```python

```