# Denison CS-181/DA-210 Homework

---

## List Comprehension Exercises

*Execute the prolog cell*

```python
import os
import io
import sys
from contextlib import redirect_stdout
from IPython.core.debugger import set_trace

def add_modules():
    """
    Starting at the current directory and proceeding up the file system
    tree, search for a directory named `modules`.  If found, and if not
    already there, add to the Python module search path.
    
    Params: None
    
    Return: None
    """
    directory = "."
    levels = 0
    while not os.path.isdir(os.path.join(directory, "modules")) and \
          levels < 5:
        directory = os.path.join(directory, "..")
        levels += 1
    module_path = os.path.abspath(os.path.join(directory, "modules"))
    if os.path.isdir(module_path):
        if not module_path in sys.path:
            sys.path.append(module_path)

add_modules()
import util
```

**Q1** Suppose we have a list of strings, named by `L1`, such as the following example:

    L1 = ['this', 'is', 'a', 'short', 'list']
    
Write code using a list comprehension to assign to `L2` a list composed of the lengths of each of the items in `L1`.

<!-- grade_id: cell-26e10e0d9c1057b9, points: 0, grade: false, solution: true -->
```python
L1 = ['this', 'is', 'a', 'short', 'list']
L2 = [len(item) for item in L1]

print(L2)
```

Output:
```
[4, 2, 1, 5, 4]
```

<!-- grade_id: cell-a97a611df60f4c22, points: 2, grade: true, solution: false -->
```python
assert L2 == [4, 2, 1, 5, 4]
```

**Q2** Abstract your work from the previous question by writing a function

    vectorLengths(data)
 
that generates and returns a new list based on the list in `data`, containing the lengths of the elements in `data`.  Use a list comprehension.

<!-- grade_id: cell-4e069cba131d2102, points: 0, grade: false, solution: true -->
```python
def vectorLengths(data):
    '''
    This function takes in a list and returns another list with
    the length of the items in each index of the original list.
    
    Params: data: The list that will be examined.
    
    Returns: A list with the lengths of the items in the
             original list.
    '''
    return [len(item) for item in data]
print(vectorLengths(['this', 'is', 'a', 'short', 'list']))
```

Output:
```
[4, 2, 1, 5, 4]
```

<!-- grade_id: cell-98dd199fb2f32ed2, points: 2, grade: true, solution: false -->
```python
assert vectorLengths(['this', 'is', 'a', 'short', 'list']) == [4, 2, 1, 5, 4]
assert vectorLengths([]) == []
assert vectorLengths(['a', 'ab', 'abc']) == [1, 2, 3]

temps = [
   ['01-03', 'Dallas', 75, 48],
   ['01-04', 'Dallas', 63, 46],
   ['01-05', 'Dallas', 71, 54]
]

assert vectorLengths(temps) == [4, 4, 4]
```

**Q3** Suppose we want to implement the binary vector operation of "less than or equal", where, given vectors `a` and `b`, a resultant vector `c` is the list of **boolean** `True` or `False` values, where each $c_i$ is the boolean comparison $a_i \le b_i$.  So, for example, if

    a = [4, 7, 2]
    b = [5, 6, 2]

the result, `c` is `[True, False, True]`.  Use a list comprehension to compute `c`, given that `a` and `b` are lists of numeric values of the same length.

<!-- grade_id: cell-debbac193f13bbb5, points: 0, grade: false, solution: true -->
```python
a = [4, 7, 2]
b = [5, 6, 2]
c = [item[0] <= item[1] for item in zip(a,b)]

print(c)
```

Output:
```
[True, False, True]
```

<!-- grade_id: cell-4e9df362c27e8bb1, points: 2, grade: true, solution: false -->
```python
assert isinstance(c, list)
assert isinstance(c[0], bool)
assert len(c) == 3
assert c == [True, False, True]
```

**Q4** Write a function

    dropTwosThrees(L)
    
that uses a list comprehension to create and return a filtered list from `L`, which only includes the elements of L whose values are *not* 2 or 3.

<!-- grade_id: cell-fd3f53043982493e, points: 0, grade: false, solution: true -->
```python
def dropTwosThrees(L):
    '''
    This function takes in a list and takes out the
    2's and 3's from the orignal list.
    
    Params: L: A list of integers.
    
    Returns: A list of integers that does not have 2's
             or 3's.
    '''
    return [item for item in L if item != 2 and item != 3]
print(dropTwosThrees([7, 2, 4, 3, 3, 5, 2, 6, 1]))
```

Output:
```
[7, 4, 5, 6, 1]
```

<!-- grade_id: cell-8abcf23c44a4ab5c, points: 2, grade: true, solution: false -->
```python
assert dropTwosThrees([7, 2, 4, 3, 3, 5, 2, 6, 1]) == [7, 4, 5, 6, 1]
```

**Q5** Write a function 

    getLoLcolumn(LoL, column_index)

that uses a list comprehension to compose and return a new list whose values are the **column** from the list of lists, `LoL` as determined by `column_index`.  So if the `LoL` were
```python
[
    ['01-07', 'Columbus', 34, 29],
    ['01-08', 'Columbus', 51, 27],
    ['01-09', 'Columbus', 47, 37],
    ['01-03', 'Dallas',   75, 48],
    ['01-04', 'Dallas',   63, 46]
]
```
and we used a column index of 3, the function would return `[29, 27, 37, 48, 46]`.

<!-- grade_id: cell-bd11e983c03c1cf0, points: 0, grade: false, solution: true -->
```python
def getLoLcolumn(LoL, column_index):
    '''
    This function returns a column of a list of lists.
    The column is returned in a list.
    
    Params: LoL: A list of lists.
            column_index: An integer that will represent
                          the column index.
    
    Returns: A list of the items in the given column spots
             of a list of lists.
    '''
    return [item[column_index] for item in LoL]
t = [
    ['01-07', 'Columbus', 34, 29],
    ['01-08', 'Columbus', 51, 27],
    ['01-09', 'Columbus', 47, 37],
    ['01-03', 'Dallas',   75, 48],
    ['01-04', 'Dallas',   63, 46]
]
print(getLoLcolumn(t, 3))
```

Output:
```
[29, 27, 37, 48, 46]
```

<!-- grade_id: cell-f9e3f921684f3df4, points: 2, grade: true, solution: false -->
```python
t = [
    ['01-07', 'Columbus', 34, 29],
    ['01-08', 'Columbus', 51, 27],
    ['01-09', 'Columbus', 47, 37],
    ['01-03', 'Dallas',   75, 48],
    ['01-04', 'Dallas',   63, 46]
]
assert getLoLcolumn(t, 3) == [29, 27, 37, 48, 46]
```

```python

```