# Denison CS-181/DA-210 Homework

---

## XPath Exercises

> The intent of this homework is to use **XPath** to obtain results, wherein the xpath string is doing the "heavy lifting" and used in a call to `.xpath()` on the root Element of an XML tree.  
>
> **If you use procedural XML instead of XPath, you will receive at most half credit per question.**
>
> For some questions, you are asked to "wrap" your xpath query into a function, or to perform some post-query step, and the previous admonition does not apply.
>
> It is perfectly fine to use an Online tool (like codebeautify) to help you develop and text your xpath string before incorporating it into its programmatic place here in this notebook.

```python
import os
import sys
import lxml

def add_modules():
    """
    Starting at the current directory and proceeding up the file system
    tree, search for a directory named `modules`.  If found, and if not
    already there, add to the Python module search path.
    
    Params: None
    
    Return: None
    """
    directory = "."
    levels = 0
    while not os.path.isdir(os.path.join(directory, "modules")) and \
          levels < 5:
        directory = os.path.join(directory, "..")
        levels += 1
    module_path = os.path.abspath(os.path.join(directory, "modules"))
    if os.path.isdir(module_path):
        if not module_path in sys.path:
            sys.path.append(module_path)

add_modules()
import util

datadir = util.resolve_dir("hierarchicaldata")
```

**Q1** Create the tree and obtain the root for `"flights.xml"` in the data directory.  Assign the root node to `froot`.  Perform any necessary imports.  By using Atom, examine this dataset and its hierarchy, noting the set of **different** tags under the root `doc`, and the structure of each.

Understanding the data will help in answering the subsequent questions.

<!-- grade_id: cell-ed96509192a96403, points: 0, grade: false, solution: true -->
```python
# Solution cell
from lxml import etree
parser = etree.XMLParser(remove_blank_text=True)

froot = etree.parse(os.path.join(datadir, "flights.xml"), parser).getroot()

```

<!-- grade_id: cell-fbd92a7e2ce334d2, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert isinstance(froot, lxml.etree._Element)
```

**Q2** Create a Python list called `passports` containing the actual passport number (i.e., not the Element/node) for each of the passengers in the tree.

<!-- grade_id: cell-eaade7c7e699e058, points: 0, grade: false, solution: true -->
```python
# Solution cell

xs_string = """/doc/Passenger/passportnumber/text()"""
passports = froot.xpath(xs_string)

passports
```

Output:
```
['123456', '123457', '000111', '000112', '000113', '000114']
```

<!-- grade_id: cell-2fcba1e02bd48ed3, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert isinstance(passports, list)
```

**Q3** Create the collection of the attribute `airId` for those airports with a tax less than 100.  Assign to Python variable `cheapairport`.

<!-- grade_id: cell-da4cb68b24779a81, points: 0, grade: false, solution: true -->
```python
# Solution cell

xs_string = """/doc/Airport[tax < "100"]/@airId"""
cheapairport = froot.xpath(xs_string)

cheapairport
```

Output:
```
['NPL', 'SPL', 'PRG', 'BDP', 'FFT']
```

<!-- grade_id: cell-6b349c3bdf1d6343, points: 2, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q5** Create a list of destination Elements for the flights where the source is NPL (North Pole).  Assign to Python variable `fromNPL`.

<!-- grade_id: cell-12a4d32fd29f0150, points: 0, grade: false, solution: true -->
```python
# Solution cell

xs_string = """/doc/Flight[source = "NPL"]/destination"""
fromNPL = froot.xpath(xs_string)

fromNPL
```

Output:
```
[<Element destination at 0x7ff40f804180>,
 <Element destination at 0x7ff40f802d80>,
 <Element destination at 0x7ff40f8020c0>,
 <Element destination at 0x7ff40f809e00>,
 <Element destination at 0x7ff40f8093c0>,
 <Element destination at 0x7ff40f809f80>,
 <Element destination at 0x7ff40f809e40>,
 <Element destination at 0x7ff40f809d00>,
 <Element destination at 0x7ff40f809b00>,
 <Element destination at 0x7ff40f809fc0>]
```

<!-- grade_id: cell-805ce96d24f64567, points: 2, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q6** Construct a list of the text contents of **all** of the children of the Flight whose flightId is LX125, assigning to `flightchildren`.

<!-- grade_id: cell-1f89ea82403fe7ab, points: 0, grade: false, solution: true -->
```python
# Solution cell

xs_string = """/doc/Flight[@flightId = "LX125"]//text()"""
flightchildren = froot.xpath(xs_string)

flightchildren
```

Output:
```
['100', '2005-12-24', '10:00:00', '11:10:00', 'LHR', 'AMS']
```

<!-- grade_id: cell-9c993e5e0c62d4e6, points: 2, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q7** Create the tree and obtain the root for `"bookstore2.xml"` in the data directory.  Assign the root node to `broot`. By using Atom, examine this dataset and its hierarchy, noting the structure of the items in the data set.

Understanding the data will help in answering the subsequent questions.

<!-- grade_id: cell-a9739dd7e61bab87, points: 0, grade: false, solution: true -->
```python
# Solution cell

broot = etree.parse(os.path.join(datadir, "bookstore2.xml"), parser).getroot()
```

<!-- grade_id: cell-e126b3beef4e19c0, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert isinstance(broot, lxml.etree._Element)
```

**Q8** Create a list of books ids named `less` that cost less than `$6`. Note that `id` is an attribute.

<!-- grade_id: cell-d615e8debfc0a057, points: 0, grade: false, solution: true -->
```python
# Solution cell
less = broot.xpath("""/catalog/book[price < "6.00"]/@id""")

less
```

Output:
```
['bk102', 'bk103', 'bk104', 'bk105', 'bk106', 'bk107', 'bk108']
```

<!-- grade_id: cell-b927e4540f49c7f8, points: 2, grade: true, solution: false -->
```python
assert True
```

**Q9** Create a list of book titles called "eva" where Eva Corets was the author. Your list `eva` should be a list of strings.

<!-- grade_id: cell-597065155a8ac10f, points: 0, grade: false, solution: true -->
```python
# Solution cell
eva = broot.xpath("""/catalog/book[author = "Corets, Eva"]/title/text()""")

eva
```

Output:
```
['Maeve Ascendant', "Oberon's Legacy", 'The Sundered Grail']
```

<!-- grade_id: cell-4172d4630a2f8ea6, points: 2, grade: true, solution: false -->
```python
assert True
```

**Q10** Find the average book price for all books that are not fantasy in this file, assigning to variable `avgprice`. **Hints** First, use XPath to get a list of the price strings (text) based on a single XPath query.  Then use a list comprehension to build a list of `float` values converting the strings to real-valued numbers.  Finally, perform the average based on the values and length of the list.

<!-- grade_id: cell-8654b21767e5a1e3, points: 0, grade: false, solution: true -->
```python
# Solution cell
avgprice = 0
prices = broot.xpath("""/catalog/book[genre != "Fantasy"]/price/text()""")
prices = [float(prices[i]) for i in range(len(prices))]
avgprice = sum(prices)/len(prices)

#I rounded the number because it was annoying me
avgprice = round(avgprice, 2)
avgprice
```

Output:
```
23.83
```

<!-- grade_id: cell-54fc268f861156e4, points: 2, grade: true, solution: false -->
```python
assert True
```

**Q11** Create a list called `lessFantasy` containing the titles of the books where the price is under `$40` and not in the fantasy genre.

<!-- grade_id: cell-2264143daec6f3bd, points: 0, grade: false, solution: true -->
```python
# Solution cell
lessFantasy = broot.xpath("""/catalog/book[genre != "Fantasy" and price < "40.00"]/title/text()""")

lessFantasy
```

Output:
```
['Lover Birds',
 'Splish Splash',
 'Creepy Crawlies',
 'Paradox Lost',
 'Microsoft .NET: The Programming Bible',
 'MSXML3: A Comprehensive Guide']
```

<!-- grade_id: cell-a8d633ddd943cc07, points: 2, grade: true, solution: false -->
```python
assert True
```

```python

```

```python

```