# Denison CS181/DA210

---

## Python File Operations

*First execute prolog cell*

```python
import os
import io
import sys
from contextlib import redirect_stdout
from IPython.core.debugger import set_trace

def add_modules():
    """
    Starting at the current directory and proceeding up the file system
    tree, search for a directory named `modules`.  If found, and if not
    already there, add to the Python module search path.
    
    Params: None
    
    Return: None
    """
    directory = "."
    levels = 0
    while not os.path.isdir(os.path.join(directory, "modules")) and \
          levels < 5:
        directory = os.path.join(directory, "..")
        levels += 1
    module_path = os.path.abspath(os.path.join(directory, "modules"))
    if os.path.isdir(module_path):
        if not module_path in sys.path:
            sys.path.append(module_path)

add_modules()
import util

datadir = util.resolve_dir("otherdata")
```

**Q1** Write a function

    wordCount(s)
    
that takes a string, `s` and, after eliminating any leading or trailing invisible/whitespace characters, divides the remaining string into whitespace separated "word" tokens and then returns the number of words.  Punctuation, in this case, is considered part of the word.

So if `s` was the string

    '\t this is a\ttest  string. \n'
    
`wordCount(s)` would return 5.

<!-- grade_id: cell-310d17340bd2a509, points: 0, grade: false, solution: true -->
```python
s = "\t this is a\ttest  string. \n"

def wordCount(s):
    """
    This function counts the words in a string and returns
    the number of words in a string.
    
    Params: s: the string that you want to know the number
            of words in it.
    
    Return: count: the number of words that are in the 
            given string.
    """
    count = 0
    sortS = s.strip()
    if sortS == "":
        return 0
    preChar = ' '
    for char in sortS:
        if (char != ' ' and char != '\a' and char != '\t' and char != '\n'):
            if (preChar == ' ' or preChar == '\a' or preChar == '\t' or preChar == '\n' ):
                count = count + 1
        preChar = char
    return count

print("count:", wordCount(s))
```

Output:
```
count: 5
```

<!-- grade_id: cell-284db49b25cba8ce, points: 2, grade: true, solution: false -->
```python
s2 = ""
assert wordCount(s2) == 0
s3 = "\n\n\n\n\n\n\n"
assert wordCount(s3) == 0

s4 = "one two three four five"
assert wordCount(s4) == 5

s1 = "\t this is a\ttest  string. \n"
assert wordCount(s1) == 5

```

**Q2** Now write a function

    fileWordCount(path)
    
that operates on the file specified by the given `path` and creates a list containing the count of words (as defined above) on each of the lines of the file.  You should **use** the function you wrote for the last question, and should not cut and paste code.  In your solution cell, demonstrate correct **invocation** of your function by composing a path to the `tennyson.txt` file in the data directory and passing that path to your function, assigning the result to `tennyson_words`.

If the `path` does not refer to an actual file, the function should return `None`.

<!-- grade_id: cell-02889d90fccfeced, points: 0, grade: false, solution: true -->
```python
def fileWordCount(path):
    """
    Makes a list that counts the amount of words on each
    line of a file.
    
    Params: path: the path of the file that will be 
            counted
    
    Return: countList: the list of words on each line of
            the given file.
            None: if the path does not lead to a file.
    """
    if os.path.isfile(path):
        fObj = open(path)
        countList = []
        for line in fObj:
            countList.append(wordCount(line))
        fObj.close()
        return countList
    return None

currDir = os.getcwd()
os.chdir("..")
os.chdir("..")
os.chdir("otherdata")
tennyson_path = os.path.join(os.getcwd(), "tennyson.txt")
os.chdir(currDir)

tennyson_words = fileWordCount(tennyson_path)
print(tennyson_words)
```

Output:
```
[4, 5, 0, 8, 3]
```

<!-- grade_id: cell-c02e0232b5456540, points: 3, grade: true, solution: false -->
```python
assert tennyson_words == [4, 5, 0, 8, 3]
assert fileWordCount("/foobar") == None

```

**Q3** Write a function

    wordLengths(s)
    
that takes a string, `s` and, after eliminating any leading or trailing invisible/whitespace characters, divides the remaining string into whitespace separated "word" tokens and composes and returns a tuple consisting of the integer lengths of each of these word-tokens.

So if `s` was the string

    '\t this is a\ttest  string. \n'
    
the leading tab-space and trailing space-newline would be dropped and the five word tokens would be `'this'`, `'is'`, `'a'`, `'test'`, `'string.'`, and the result would be the tuple

    (4, 2, 1, 4, 7)

<!-- grade_id: cell-fcca7e742b3ffafb, points: 0, grade: false, solution: true -->
```python
s = "\t this is a\ttest  string. \n"
def wordLengths(s):
    """
    Makes a tuple that counts the amount of characters 
    in each word of a string.
    
    Params: s: the string that is going to have the
            words counted.
    
    Return: wrdCountTup: a tuple that contains the 
            amount of leters in each word of a given
            string.
    """
    count = 0
    sortS = s.strip()
    wrdCountTup = ()
    preChar = ''
    if sortS == "":
        return wrdCountTup
    for char in sortS:
        if (char != ' ' and char != '\a' and char != '\t' and char != '\n'):
            count = count + 1
        elif (preChar != ' ' and preChar != '\a' and preChar != '\t' and preChar != '\n'):
            wrdCountTup = wrdCountTup + (count,)
            count = 0
        preChar = char
    wrdCountTup = wrdCountTup + (count,)
    return wrdCountTup

print(wordLengths(s))
```

Output:
```
(4, 2, 1, 4, 7)
```

<!-- grade_id: cell-fa42e337ba7afd5a, points: 2, grade: true, solution: false -->
```python
s2 = ""
assert wordLengths(s2) == ()
s3 = "\n\n\n\n\n\n\n"
assert wordLengths(s3) == ()

s4 = "one two three four five"
assert wordLengths(s4) == (3, 3, 5, 4, 4)

s = "\t this is a\ttest  string. \n"
assert wordLengths(s) == (4, 2, 1, 4, 7)
```

**Q4** Now write a function

    fileWordLengths(path)
    
that operates on the file specified by the given `path` and creates a list containing, for each line, a tuple with the word lengths in the corresponding line.  In your solution cell, demonstrate correct **invocation** of your function by composing a path to the `tennyson.txt` file in the data directory and passing that path to your function, assigning the result to `tennyson_wordlengths`.

If the `path` does not refer to an actual file, the function should return `None`.

<!-- grade_id: cell-dce165fd47e3120e, points: 0, grade: false, solution: true -->
```python
def fileWordLengths(path):
    """
    This function counts the amount of letters in each word 
    on each line of a given file and returns a list of 
    tuples. The tuples have the amount of letters in each
    word and each item in the list is each line in the file.
    
    Params: path: the path of the file that will be counted.
    
    Return: countList: the list of tuples that contain the
            amount of letters for each word in each line of
            the given file.
            None: if the path given does not lead to a file.
    """
    if os.path.isfile(path):
        fObj = open(path)
        countList = []
        for line in fObj:
            countList.append(wordLengths(line))
        fObj.close()
        return countList
    return None

tennyson_wordlengths = fileWordLengths(tennyson_path)
print(tennyson_wordlengths)
```

Output:
```
[(8, 3, 7, 5), (3, 5, 4, 3, 5), (), (3, 3, 5, 2, 2, 7, 2, 9), (4, 1, 7)]
```

<!-- grade_id: cell-36d5ec6ed4850aa7, points: 3, grade: true, solution: false -->
```python
assert tennyson_wordlengths == [(8, 3, 7, 5), (3, 5, 4, 3, 5), (), (3, 3, 5, 2, 2, 7, 2, 9), (4, 1, 7)]

```

The next two questions involve the `temps.csv` file in the data directory.  It's contents are reflected here:

```
01-03, Columbus, 47, 13
01-04, Columbus, 48, 31
01-05, Columbus, 36, 23
01-06, Columbus, 50, 32
01-07, Columbus, 34, 29
01-08, Columbus, 51, 27
01-09, Columbus, 47, 37
01-03,   Dallas, 75, 48
01-04,   Dallas, 63, 46
01-05,   Dallas, 71, 54
01-06,   Dallas, 62, 39
01-07,   Dallas, 42, 39
01-08,   Dallas, 41, 37
01-09,   Dallas, 56, 36
```
where fields are separated by commas; spaces are embedded within for readability, there is **no** header line, and the fields represent the columns of *date*, *city*, *hightemp*, and *lowtemp*.

**Q5** Write code to process this file and generate **four parallel lists**, named `dates`, `cities`, `hightemps` and `lowtemps`.  **Make sure the temperature lists have integer elements**, and the date and city lists have string elements.  You are required to use file operations for this processing, and may not employ other packages.

<!-- grade_id: cell-98c72a122d657922, points: 0, grade: false, solution: true -->
```python
currDir = os.getcwd()
os.chdir("..")
os.chdir("..")
os.chdir("otherdata")
temps_path = os.path.join(os.getcwd(), "temps.csv")
os.chdir(currDir)

fObj = open(temps_path)
dates = []
cities = []
hightemps = []
lowtemps = []
for line in fObj:
    trimmed = line.strip()
    fields = trimmed.split(',')
    dates.append(fields[0])
    cities.append(fields[1].strip())
    hightemps.append(int(fields[2]))
    lowtemps.append(int(fields[3]))
fObj.close()
```

<!-- grade_id: cell-8262a676135d3021, points: 3, grade: true, solution: false -->
```python
assert dates == ['01-03', '01-04','01-05','01-06','01-07','01-08','01-09'] * 2
assert cities == ['Columbus'] * 7 + ['Dallas'] * 7
assert hightemps == [47, 48, 36, 50, 34, 51, 47, 75, 63, 71, 62, 42, 41, 56]
assert lowtemps == [13, 31, 23, 32, 29, 27, 37, 48, 46, 54, 39, 39, 37, 36]
```

**Q6** Repeat the processing of `temps.csv` generating a single **dictionary** representing the data.  The dictionary should have exactly one mapping per line of the file.  A dictionary item should map a key which is a **tuple** containing the date and the city (the independent variables in this data set), and this key maps to a tuple containing the hightemp and lowtemp.  

So, for example, the first entry added to the dictionary will map `('02-03', 'Columbus')` to `(47, 13)`.

Name the dictionary `temps`.  You are required to use file operations for this processing, and may not employ other packages.

<!-- grade_id: cell-28ddcb5593f8c015, points: 0, grade: false, solution: true -->
```python
temps = {}
for i in range(len(cities)):
    temps[(dates[i], cities[i])] = (hightemps[i], lowtemps[i])
    
```

<!-- grade_id: cell-7dc02b59c96dde63, points: 3, grade: true, solution: false -->
```python
assert temps == {('01-03', 'Columbus'): (47, 13),
 ('01-04', 'Columbus'): (48, 31),
 ('01-05', 'Columbus'): (36, 23),
 ('01-06', 'Columbus'): (50, 32),
 ('01-07', 'Columbus'): (34, 29),
 ('01-08', 'Columbus'): (51, 27),
 ('01-09', 'Columbus'): (47, 37),
 ('01-03', 'Dallas'): (75, 48),
 ('01-04', 'Dallas'): (63, 46),
 ('01-05', 'Dallas'): (71, 54),
 ('01-06', 'Dallas'): (62, 39),
 ('01-07', 'Dallas'): (42, 39),
 ('01-08', 'Dallas'): (41, 37),
 ('01-09', 'Dallas'): (56, 36)}
```

```python

```