# Denison CS-181/DA-210 Homework

---

## XPath Flipped Class Notebook

```python
import os
import json
import sys
from lxml import etree

def add_modules():
    """
    Starting at the current directory and proceeding up the file system
    tree, search for a directory named `modules`.  If found, and if not
    already there, add to the Python module search path.
    
    Params: None
    
    Return: None
    """
    directory = "."
    levels = 0
    while not os.path.isdir(os.path.join(directory, "modules")) and \
          levels < 5:
        directory = os.path.join(directory, "..")
        levels += 1
    module_path = os.path.abspath(os.path.join(directory, "modules"))
    if os.path.isdir(module_path):
        if not module_path in sys.path:
            sys.path.append(module_path)

add_modules()
import util

datadir = util.resolve_dir("hierarchicaldata")
```

### References

- [CodeBeautify Online XPath](https://codebeautify.org/Xpath-Tester)
- [XPather Online XPath](http://xpather.com)

**Q1** Write the code to
- Create a parser, `myparser` that can parse XML and eliminate blank whitespace,
- Build a path to the `"bookstore2.xml" file in the hierarchical data directory,
- Assign to `booktree` the result of parsing the file (an ElementTree) using `myparser`,
- Assign to `bookroot` the root of the tree

<!-- grade_id: cell-ba4ebfbddac5bbaf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
myparser = etree.XMLParser(remove_blank_text=True)
path = os.path.join(datadir, "bookstore2.xml")
booktree = etree.parse(path, myparser)
bookroot = booktree.getroot()
type(booktree)
```

Output:
```
lxml.etree._ElementTree
```

<!-- grade_id: cell-75e10f1b4829c2cb, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q2** Find the set of nodes representing magazines in the data set. 


<!-- grade_id: cell-357336b6af4f857f, points: 0, grade: false, solution: true -->
```python
xp_string = """/catalog/magazine"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'/catalog/magazine'
Length of result:  4
[<Element magazine at 0x7fdb0bd97a80>,
 <Element magazine at 0x7fdb0bd971c0>,
 <Element magazine at 0x7fdb0bd97a00>,
 <Element magazine at 0x7fdb0bd97a40>]
```

<!-- grade_id: cell-96384cd735a2c883, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q3** Find the textual title of magazines. 

<!-- grade_id: cell-1ac52e7570e9d253, points: 0, grade: false, solution: true -->
```python
xp_string = """/catalog/magazine/title/text()"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'/catalog/magazine/title/text()'
Length of result:  4
['National Geographic',
 'National Geographic',
 'Newsweek',
 "Hector and Jeff's Database Hints"]
```

<!-- grade_id: cell-922e5609acda27fb, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q4** Find the set of years of magazines. 


<!-- grade_id: cell-b3b8ca78e91b9071, points: 0, grade: false, solution: true -->
```python
xp_string = """/catalog/magazine/@year"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'/catalog/magazine/@year'
Length of result:  4
['2021', '2020', '2021', '2009']
```

<!-- grade_id: cell-416a61dab920a144, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q5** Find the node for the book whose id is `bk101`



<!-- grade_id: cell-f1738b70e59140a2, points: 0, grade: false, solution: true -->
```python
xp_string = """/catalog/book[@id="bk101"]"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'/catalog/book[@id="bk101"]'
Length of result:  1
[<Element book at 0x7fdb0bda8b40>]
```

<!-- grade_id: cell-e75b6dead0f17c0e, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q6** Find all the children for the book whose id is `bk101


<!-- grade_id: cell-e062a079743d6fa3, points: 0, grade: false, solution: true -->
```python
xp_string = """/catalog/book[@id="bk101"]/*"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'/catalog/book[@id="bk101"]/*'
Length of result:  6
[<Element author at 0x7fdb0bdadfc0>,
 <Element title at 0x7fdb0bdb0280>,
 <Element genre at 0x7fdb0bdb0140>,
 <Element price at 0x7fdb0bdb0200>,
 <Element publish_date at 0x7fdb0bdb0240>,
 <Element description at 0x7fdb0bdb0300>]
```

<!-- grade_id: cell-7146bb6042c8c2b0, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q7** Find the set of nodes for book elements in the data set.


<!-- grade_id: cell-d526dae005c76a1f, points: 0, grade: false, solution: true -->
```python
xp_string = """/catalog/book"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'/catalog/book'
Length of result:  12
[<Element book at 0x7fdb0bda8b40>,
 <Element book at 0x7fdb0bdb6740>,
 <Element book at 0x7fdb0bdb8840>,
 <Element book at 0x7fdb0bdb87c0>,
 <Element book at 0x7fdb0bdb8800>,
 <Element book at 0x7fdb0bdb88c0>,
 <Element book at 0x7fdb0bdb8900>,
 <Element book at 0x7fdb0bdb8940>,
 <Element book at 0x7fdb0bdb8980>,
 <Element book at 0x7fdb0bdb8880>,
 <Element book at 0x7fdb0bdb89c0>,
 <Element book at 0x7fdb0bdb8a00>]
```

<!-- grade_id: cell-07deb70db6608949, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q8** Find the set of nodes for the author of books in the data set.


<!-- grade_id: cell-63c22a12d1f9f8f1, points: 0, grade: false, solution: true -->
```python
xp_string = """/catalog/book/author"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'/catalog/book/author'
Length of result:  12
[<Element author at 0x7fdb0bdadfc0>,
 <Element author at 0x7fdb0bdb7580>,
 <Element author at 0x7fdb0bdbc540>,
 <Element author at 0x7fdb0bdbc3c0>,
 <Element author at 0x7fdb0bdbc400>,
 <Element author at 0x7fdb0bdbc500>,
 <Element author at 0x7fdb0bdb7b40>,
 <Element author at 0x7fdb0bdb7b00>,
 <Element author at 0x7fdb0bdb7ac0>,
 <Element author at 0x7fdb0bdb7b80>,
 <Element author at 0x7fdb0bdb7900>,
 <Element author at 0x7fdb0bdb78c0>]
```

<!-- grade_id: cell-2ffe450109961c01, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q9** Find the set of ids of books. 


<!-- grade_id: cell-1351f1928d5b43d1, points: 0, grade: false, solution: true -->
```python
xp_string = """/catalog/book/@id"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'/catalog/book/@id'
Length of result:  12
['bk101',
 'bk102',
 'bk103',
 'bk104',
 'bk105',
 'bk106',
 'bk107',
 'bk108',
 'bk109',
 'bk110',
 'bk111',
 'bk112']
```

<!-- grade_id: cell-f644d2d0ac46fcaa, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q10** Find the node for the magazine whose id is `mag102`


<!-- grade_id: cell-ced5938a313c9bb8, points: 0, grade: false, solution: true -->
```python
xp_string = """/catalog/magazine[@id="mag102"]"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'/catalog/magazine[@id="mag102"]'
Length of result:  1
[<Element magazine at 0x7fdb0bd971c0>]
```

<!-- grade_id: cell-800f1cca9ed067cf, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q11** Find the text of the title for the magazine whose id is `mag102`


<!-- grade_id: cell-6947dae27745936e, points: 0, grade: false, solution: true -->
```python
xp_string = """/catalog/magazine[@id="mag102"]/title/text()"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'/catalog/magazine[@id="mag102"]/title/text()'
Length of result:  1
['National Geographic']
```

<!-- grade_id: cell-ec82cdc5b0290066, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q12** Find the text of the title child for any child of catalog.  Use an absolute path and a wildcard.


<!-- grade_id: cell-40de825d393fe5ce, points: 0, grade: false, solution: true -->
```python
xp_string = """//catalog/*/title/text()"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'//catalog/*/title/text()'
Length of result:  16
["XML Developer's Guide",
 'Midnight Rain',
 'National Geographic',
 'National Geographic',
 'Maeve Ascendant',
 "Oberon's Legacy",
 'The Sundered Grail',
 'Lover Birds',
 'Splish Splash',
 'Creepy Crawlies',
 'Paradox Lost',
 'Microsoft .NET: The Programming Bible',
 'MSXML3: A Comprehensive Guide',
 'Visual Studio 7: A Comprehensive Guide',
 'Newsweek',
 "Hector and Jeff's Database Hints"]
```

<!-- grade_id: cell-61b1db97a9d181b5, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q13** Find the text of any title element in the entire tree. 


<!-- grade_id: cell-5c260f60b51f54eb, points: 0, grade: false, solution: true -->
```python
xp_string = """//catalog/*/*/text()"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'//catalog/*/*/text()'
Length of result:  76
['Gambardella, Matthew',
 "XML Developer's Guide",
 'Computer',
 '44.95',
 '2000-10-01',
 'An in-depth look at creating applications\n      with XML.',
 'Ralls, Kim',
 'Midnight Rain',
 'Fantasy',
 '5.95',
 '2000-12-16',
 'A former architect battles corporate zombies,\n      an evil sorceress, and her own childhood to become queen\n      of the world.',
 'National Geographic',
 'National Geographic',
 'Corets, Eva',
 'Maeve Ascendant',
 'Fantasy',
 '5.95',
 '2000-11-17',
 'After the collapse of a nanotechnology\n      society in England, the young survivors lay the\n      foundation for a new society.',
 'Corets, Eva',
 "Oberon's Legacy",
 'Fantasy',
 '5.95',
 '2001-03-10',
 'In post-apocalypse England, the mysterious\n      agent known only as Oberon helps to create a new life\n      for the inhabitants of London. Sequel to Maeve\n      Ascendant.',
 'Corets, Eva',
 'The Sundered Grail',
 'Fantasy',
 '5.95',
 '2001-09-10',
 "The two daughte
```

<!-- grade_id: cell-a88dfe512877c389, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q14** Find the value of all `id` attributes anywhere in the tree.

<!-- grade_id: cell-dc612eb6883c55ed, points: 0, grade: false, solution: true -->
```python
xp_string = """/catalog/*/@id"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'/catalog/*/@id'
Length of result:  16
['bk101',
 'bk102',
 'mag101',
 'mag102',
 'bk103',
 'bk104',
 'bk105',
 'bk106',
 'bk107',
 'bk108',
 'bk109',
 'bk110',
 'bk111',
 'bk112',
 'mag103',
 'mag104']
```

<!-- grade_id: cell-5439f5e8ce1c5bec, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q15** Find the value of **all** attributes.  

<!-- grade_id: cell-a2422d6ce7a0db16, points: 0, grade: false, solution: true -->
```python
xp_string = """/catalog/*/@*"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'/catalog/*/@*'
Length of result:  23
['bk101',
 'bk102',
 'mag101',
 'January',
 '2021',
 'mag102',
 'February',
 '2020',
 'bk103',
 'bk104',
 'bk105',
 'bk106',
 'bk107',
 'bk108',
 'bk109',
 'bk110',
 'bk111',
 'bk112',
 'mag103',
 'February',
 '2021',
 'mag104',
 '2009']
```

<!-- grade_id: cell-f0f7c6e141c3c287, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q16** Find the set of magazine nodes whose month is February.

<!-- grade_id: cell-d84bfcf3c57e9c1a, points: 0, grade: false, solution: true -->
```python
xp_string = """/catalog//magazine[@month="February"]"""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'/catalog//magazine[@month="February"]'
Length of result:  2
[<Element magazine at 0x7fdb0bd971c0>, <Element magazine at 0x7fdb0bd97a00>]
```

<!-- grade_id: cell-4a39e79465681107, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q17** Find the parent(s) of the set of magazine nodes whose month is February.

<!-- grade_id: cell-12dbb4cc19cbb529, points: 0, grade: false, solution: true -->
```python
xp_string = """/catalog//magazine[@month="February"]/.."""
print(repr(xp_string))
result = bookroot.xpath(xp_string)
print("Length of result: ", len(result))
result
```

Output:
```
'/catalog//magazine[@month="February"]/..'
Length of result:  1
[<Element catalog at 0x7fdb0bd87d40>]
```

<!-- grade_id: cell-d9b9b73edeccdb23, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert True
```

```python

```