# Denison CS-181/DA-210 Homework

---

## Regular Expression Programming in Python

*Execute the prolog cell*

```python
import re

import os
import io
import sys

from contextlib import redirect_stdout
from IPython.core.debugger import set_trace

def add_modules():
    """
    Starting at the current directory and proceeding up the file system
    tree, search for a directory named `modules`.  If found, and if not
    already there, add to the Python module search path.
    
    Params: None
    
    Return: None
    """
    directory = "."
    levels = 0
    while not os.path.isdir(os.path.join(directory, "modules")) and \
          levels < 5:
        directory = os.path.join(directory, "..")
        levels += 1
    module_path = os.path.abspath(os.path.join(directory, "modules"))
    if os.path.isdir(module_path):
        if not module_path in sys.path:
            sys.path.append(module_path)

add_modules()
import util
```

**Q** From the reading in the book and from any class discussion, explain why we use **raw** strings when we specify regular expression patterns.  In your discussion, be clear on why we do **not** use raw strings on the target strings within which we are looking for matches to a pattern.

<!-- grade_id: cell-3db8e98ebfa3c830, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q** Suppose we want the **regular expression** pattern of:
    
`\w+\s+\\\.`

Using a **regular Python string** constant literal, assign to `p` the string that results in this pattern.

<!-- grade_id: cell-a2e514b33f6dd18b, points: 0, grade: false, solution: true -->
```python
p = "\\w+\\s+\\\\\\."

print(p)
```

Output:
```
\w+\s+\\\.
```

<!-- grade_id: cell-a7ebd8a50d6f6319, points: 1, grade: true, solution: false -->
```python
assert True
```

**Q** Suppose we again want the **regular expression** pattern of:
    
`\w+\s+\\\.`

Using a **raw Python string** constant literal, assign to `p` the string that results in this pattern.

<!-- grade_id: cell-15e729adc47ef66a, points: 0, grade: false, solution: true -->
```python
p = r"\w+\s+\\\."

print(p)
```

Output:
```
\w+\s+\\\.
```

<!-- grade_id: cell-8d311ce6538120fc, points: 1, grade: true, solution: false -->
```python
assert True
```

#### Primary Functions in `re`

Function | Return Type | Short Description
:-------:|-------------|-----------------------------------------------------------------
`search()` | *Match Object* | Find first match of pattern *anywhere* in target
`match()` | *Match Object* | Find match of pattern at *start* of target
`findall()` | *List of strings* | Find all matches and return list of matched strings, or list of tuples of captured groups
`finditer()` | *Match Object Iterator* | Find all matches through an iterator of successful match objects

```python
target = """From: Robert Elz <kre@munnari.OZ.AU>
To: Chris Garrigues <cwg-dated-1030376441.95dd45@DeepEddy.com>
Cc: exmh-workers@example.com
Subject: Re: New Sequences Window From: external
Sender: exmh-workers-admin@example.com
Date: Thu, 22 Aug 2002 21:55:18 +0700
"""
```

**Q** Find the first match of the `.com` pattern within the given text. Use print to show a representation of the match object returned.

<!-- grade_id: cell-f753195de910fd77, points: 1, grade: true, solution: true -->
```python
pattern = r"\.com"
print(re.search(pattern,target))
```

Output:
```
<re.Match object; span=(94, 98), match='.com'>
```

**Q** Looking carefully at `target`, how do you interpret the meaning of the `span` part of the match object?

<!-- grade_id: cell-b10b29c5cda50206, points: 1, grade: true, solution: true -->
Span represents the starting character and the character right after that your pattern is located at in the taget. So the pattern spans those characters in the target.

**Q** Write a regular expression whose pattern matches the literal `From` followed by a colon and a space followed by the value of the From field up to the end of the line.

Then invoke on `target`, using a function that only succeeds if the pattern is matchted at the start of the target.

<!-- grade_id: cell-b8db7ec5167443af, points: 1, grade: true, solution: true -->
```python
pattern = r"From:+."
re.match(pattern,target)
```

Output:
```
<re.Match object; span=(0, 6), match='From: '>
```

**Q** Generate a list of strings giving *all* the matches of the From pattern in the same text. Print the resultant list of strings.

<!-- grade_id: cell-694bda65865b5519, points: 1, grade: true, solution: true -->
```python
print(re.findall(pattern,target))
```

Output:
```
['From: ', 'From: ']
```

**Q** Use a loop to find and then print all matches (the string representation of the corresponding match object) using the same pattern on the same target.

<!-- grade_id: cell-ee93aed5b3f23d2d, points: 1, grade: true, solution: true -->
```python
place = 0
for i in range(len(re.findall(pattern,target))):
    m = re.search(pattern,target[place:])
    place = m.end()
    print(m)
    
```

Output:
```
<re.Match object; span=(0, 6), match='From: '>
<re.Match object; span=(157, 163), match='From: '>
```

Method | Description
--------|-----------------
`Match.group(i)` | Return the string for the $i$th capture group.  For example, if the match object is `m`, then `m.group(2)` gives the string of capture group 2.  Match objects define an accessor shorthand, so `m[2]` does the same thing.  Group 0 is the full match.
`Match.start(i)` | Return the index in target of the start of the string of the $i$th capture group.
`Match.end(i)` | Return the index in target of the end of the string of the $i$th capture group.
`Match.span(i)` | Return the tuple *span* of the $i$th capture group, where the span has both the *start* and the *up to* index values of the match within the string.


**Q** In the following, we give the target and the pattern.  Write a loop that iterates over the match objects and, within the loop, build a tuple with three elements:

- the string value of the first capture group (the person named in the target)
- the string value of the second capture group (the fruit)
- a float, converted from the third capture group (the price)

Print the resultant tuple at each iteration.

<!-- grade_id: cell-5f9ba8acd3fbae2b, points: 1, grade: true, solution: true -->
```python
target = """Alice says the price of oranges is $2.50
Bob says the price of apples is $1.75
Tom thinks the price of tangerines is $2.25
"""

pattern = r'(\w+).*of\s*(\w+).*\$(\d+\.\d{2})'
for m in re.finditer(pattern,target):
    print((m.group(1), m.group(2), float(m.group(3))))
    
```

Output:
```
('Alice', 'oranges', 2.5)
('Bob', 'apples', 1.75)
('Tom', 'tangerines', 2.25)
```

**Q**  Write a  a general matching function

    getMatches(pattern, text, flags=0)
   
that find as many matches as possible in the given `text`, passing `flags` as appropriate to modify the means of matching. The function returns a list with an entry for each found match.  

For a given match, the entry is a list of tuples, which we call a "capture list."  The tuples are the string and index for each of the capture groups found within the match. Hence, the function returns a list of lists of tuples.

<!-- grade_id: cell-45ed3d665a7c90c0, points: 1, grade: true, solution: true -->
```python
def getMatches(pattern, text, flags=0):
    
```