# Denison CS181/DA210

---

## Instructions and Preliminary Cells

For many homework notebooks, there well be a prolog that contains a code cell containing setup code, like imports, global variable assignments, and utility functions.  You should execute this code cell prior to writing any solution code.

"Execution" of a markdown cell causes it to be rendered, and clicking on the cell and hitting the Return key allows you to edit the cell.  Execution of a code cell causes the statements in the cell to be run by the Python interpreter.  These cells can be edited by clicking on the cell and typing in the new infornation.

Execution of a single cell (markdown or code) can be accomplished with the keyboard by holding down the Shift key and then hitting the Return key.  This is much faster than using a mouse back and forth and clicking the Arrow/Play button in the toolbar.

---

```python
import os
import io
import sys
from contextlib import redirect_stdout

def add_modules():
    """
    Starting at the current directory and proceeding up the file system
    tree, search for a directory named `modules`.  If found, and if not
    already there, add to the Python module search path.
    
    Params: None
    
    Return: None
    """
    directory = "."
    levels = 0
    while not os.path.isdir(os.path.join(directory, "modules")) and \
          levels < 5:
        directory = os.path.join(directory, "..")
        levels += 1
    module_path = os.path.abspath(os.path.join(directory, "modules"))
    if os.path.isdir(module_path):
        if not module_path in sys.path:
            sys.path.append(module_path)

add_modules()
import util

datadir = util.resolve_dir("publicdata")
```

## Problem Structure

After any prolog, the typical homework notebook will proceed with a set of problems.  Each problem is stated in a markdown cell.  You will not make any changes to the problem statement cell.

What follows the problem statement is generally one of the following:

1. If the question requires you to answer in English, there will be a following markdown cell, where you can enter (and format) your answer.  These answers will always be manually graded.

2. If the question requires you to write Python code, there will be a following code cell.  Within the code cell, you will give your Python answer.  Sometimes this code cell will give you some preliminary code and/or some code that somehow uses or incorporates your answer and helps you see what is expected.  Whether these is initial or following code, the cell will have the two line sequence:

```python
    # YOUR CODE HERE
    raise NotImplementedError()
```

You will **delete both of these lines** and replace them with your solution to the problem.  If you do not delete the second line, then even if you have correct code, execution of the cell will result in an error and the question will be assessed as incorrect.

There are three main types of Python coding questions:

A. We ask you to use `print` to demonstrate your answer.

B. We ask you to write code and assign one or more answers to one or more Python variables, which we explicitly name.

C. We ask you to write a Python function that performs some specified computation, uses parameters as specified, and *returns* some result.

For code-type questions, there will often be an additional following code cell.  The purpose of this code in some way "checks" your code.  The purpose is (mostly) twofold:

1. To give you some (preliminary) feedback on your answer.
2. To permit some automation of checking your homework submissions by the course staff.

The way the checking cell works is to, in some way, evaluate your code and compare the answer that you produce to the answer that is expected.  If your answer matches the expected answer, the cell should appear to simply execute (with no output).  But if your answer does not match the expected answer, a Python `exception` is raised, displaying some type of error message.

To use this mechanism for good and productive feedback, you should look carefully at:

- the error message, which will give specifics on the mismatch,
- the **point** at which the checking cell resulted in an error.  You may find that your code is partially right, and "passed" some of the tests, but not others.
- **what your code produced** and **what the expected result was**.  You can (and should) use good debugging techniques to figure out how and why the deviation occurred.

I should note that sometimes giving checking code that is visible to the student can "give away" the answer.  So the system we use allows for "hidden tests" that are not visible in your Notebook, but get added and checked when we use the automated system for assessing your assignment.

---

## Example Problems

In this section, we will present some example problems, with a variety of types of problems.

**Q1** (Markdown Solution) Give me a numeric list of your top three movies.

<!-- grade_id: cell-5cd29a56eedf32df, points: 1, grade: true, solution: true -->
1. Forest Gump
2. Deadpool
3. How to Train Your Dragon

**Q2** (Code Solution: setting variable(s)) We define variable `F` with a real-valued temperature in degrees Farenheit.  Write code to compute the equivalent Celsius temperature, and assign to `result`.

<!-- grade_id: cell-ab5892a5ada0c7a8, points: 0, grade: false, solution: true -->
```python
# Solution cell

F = 212.0

result = (F - 32) * 5 / 9

print("Result:", result)
```

Output:
```
Result: 100.0
```

<!-- grade_id: cell-7f055bf7d6222e19, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert result == 100.0
```

**Q3** (Code Solution: printing) Write a loop to print the character sequence `Hello World!` five times, each on a separate line.  Do this within the `with` construct.

<!-- grade_id: cell-134b3db88732a88d, points: 0, grade: false, solution: true -->
```python
result = io.StringIO()
with redirect_stdout(result):
    for i in range(5):
        print('Hello World!')
print(result.getvalue())
```

Output:
```
Hello World!
Hello World!
Hello World!
Hello World!
Hello World!
```

<!-- grade_id: cell-7edee43a8bd7e532, points: 1, grade: true, solution: false -->
```python
# Testing Cell
s = result.getvalue()
lines = s.split('\n')

assert s.count('\n') == 5
assert s[-1] == '\n'
for line in lines[:-1]:
    assert line == "Hello World!"
```

**Q4** (Code Solution: function) Suppose we want a recursive function

    recursive_sum(mylist)

to sum the value of the elements in a list.  We give you the base case, and want you to write the recursive case for this function.

<!-- grade_id: cell-29e406974bda906d, points: 0, grade: false, solution: true -->
```python
# Solution Cell

def recursive_sum(mylist):
    """
    Sum the values of a list recursively.
    
    Params:
      mylist: the list of values, assumed to be numeric
      
    Return:
      integer or real-valued sum of the values
    """
    
    # Base case(s)
    if len(mylist) == 0:
        return 0
    if len(mylist) == 1:
        return mylist[0]
    firstvalue = mylist[0]
    rest_value = recursive_sum(mylist[1:])
    return firstvalue + rest_value

result = recursive_sum([1, 2, 3])
print(result)
```

Output:
```
6
```

<!-- grade_id: cell-eeacc5b2cba8e84e, points: 1, grade: true, solution: false -->
```python
assert recursive_sum([]) == 0
assert recursive_sum([42.1]) == 42.1
assert recursive_sum([1, 2, 3]) == 6
```

```python

```