# Denison CS-181/DA-210 Homework

---

## GroupBy Exercises

```python
import os
import io
import sys
import re
import pandas as pd

from contextlib import redirect_stdout

def add_modules():
    """
    Starting at the current directory and proceeding up the file system
    tree, search for a directory named `modules`.  If found, and if not
    already there, add to the Python module search path.
    
    Params: None
    
    Return: None
    """
    directory = "."
    levels = 0
    while not os.path.isdir(os.path.join(directory, "modules")) and \
          levels < 5:
        directory = os.path.join(directory, "..")
        levels += 1
    module_path = os.path.abspath(os.path.join(directory, "modules"))
    if os.path.isdir(module_path):
        if not module_path in sys.path:
            sys.path.append(module_path)

add_modules()
import util

datadir = util.resolve_dir("tabulardata")
```

**Q1** Read the CSV file, "indicators4.csv" from the tabular data directory into a data frame named `indicators`. Use index columns of `year` and `code` (in that order).  Print the prefix of the first 8 rows of the data frame.

```python
# Student Experimentation Cell


```

<!-- grade_id: cell-5a5b2c18f07ac2c2, points: 0, grade: false, solution: true -->
```python
# Solution Cell

result = io.StringIO()
with redirect_stdout(result):
    indicators = pd.read_csv(os.path.join(datadir, "indicators4.csv"), index_col = ['year', 'code'])
    print(indicators.head(8))
print(result.getvalue())
```

Output:
```
country       land               income               region  \
year code                                                                       
2015 CHN           China  9388210.0  Upper middle income  East Asia & Pacific   
     IND           India  2973190.0  Lower middle income           South Asia   
     USA   United States  9147420.0          High income        North America   
     VNM         Vietnam   310070.0  Lower middle income  East Asia & Pacific   
2016 CHN           China  9388210.0  Upper middle income  East Asia & Pacific   
     IND           India  2973190.0  Lower middle income           South Asia   
     USA   United States  9147420.0          High income        North America   
     VNM         Vietnam   310070.0  Lower middle income  East Asia & Pacific   

               pop       gdp   life     cell  
year code                                     
2015 CHN   1371.22  11015.54  76.09  1291.98  
     IND   1310.15   2103.59  68.30  1001.06  
     USA    320.74
```

<!-- grade_id: cell-a00aad80d673b505, points: 1, grade: true, solution: false -->
```python
# Testing Cell

assert isinstance(indicators, pd.core.frame.DataFrame)
assert result.getvalue() == '                 country       land               income               region  \\\nyear code                                                                       \n2015 CHN           China  9388210.0  Upper middle income  East Asia & Pacific   \n     IND           India  2973190.0  Lower middle income           South Asia   \n     USA   United States  9147420.0          High income        North America   \n     VNM         Vietnam   310070.0  Lower middle income  East Asia & Pacific   \n2016 CHN           China  9388210.0  Upper middle income  East Asia & Pacific   \n     IND           India  2973190.0  Lower middle income           South Asia   \n     USA   United States  9147420.0          High income        North America   \n     VNM         Vietnam   310070.0  Lower middle income  East Asia & Pacific   \n\n               pop       gdp   life     cell  \nyear code                                     \n2015 CHN   1371.22  11015.54  76.09  1291.98  \n     IND   1310.15   2103.59  68.30  1001.06  \n     USA    320.74  18219.30  78.69   382.31  \n     VNM     92.68    193.24  76.05   120.32  \n2016 CHN   1378.66  11137.95  76.25  1364.93  \n     IND   1324.51   2290.43  68.56  1127.81  \n     USA    323.07  18707.19  78.54   387.40  \n     VNM     93.64    205.28  76.25   120.60  \n'

```

**Q2** Find the min and max of variables `pop` and `gdp` **for each year**, printing the result.

```python
# Student Experimentation Cell


```

<!-- grade_id: cell-87dc826bdc157e84, points: 0, grade: false, solution: true -->
```python
# Solution Cell

result = io.StringIO()
with redirect_stdout(result):
    print(indicators.groupby('year').agg({'pop':['min','max'],'gdp':['min','max']}))
print(result.getvalue())
```

Output:
```
pop              gdp          
        min      max     min       max
year                                  
2015  92.68  1371.22  193.24  18219.30
2016  93.64  1378.66  205.28  18707.19
2017  94.60  1386.40  223.78  19485.39
```

<!-- grade_id: cell-3eb285cbb656fd31, points: 2, grade: true, solution: false -->
```python
# Testing Cell

assert result.getvalue() == '        pop              gdp          \n        min      max     min       max\nyear                                  \n2015  92.68  1371.22  193.24  18219.30\n2016  93.64  1378.66  205.28  18707.19\n2017  94.60  1386.40  223.78  19485.39\n'

```

**Q3** Find the min and max of variables `pop` and `gdp` for each country code, printing the result.

```python
# Student Experimentation Cell


```

<!-- grade_id: cell-1e62db1ab9a3a2af, points: 0, grade: false, solution: true -->
```python
# Solution Cell

result = io.StringIO()
with redirect_stdout(result):
    print(indicators.groupby('code').agg({'pop':['min','max'],'gdp':['min','max']}))
print(result.getvalue())
```

Output:
```
pop                gdp          
          min      max       min       max
code                                      
CHN   1371.22  1386.40  11015.54  12143.49
IND   1310.15  1338.66   2103.59   2652.55
USA    320.74   325.15  18219.30  19485.39
VNM     92.68    94.60    193.24    223.78
```

<!-- grade_id: cell-fd2134651e589bab, points: 1, grade: true, solution: false -->
```python
# Testing Cell

assert True

```

**Q4** Suppose we want to compare countries in income categories.  To do this, we first need to summarize variables by country, since there are multiple years per country.  Given the resultant dataframe of stats by country, we can then group by and aggregate by income category.

1. Compute the per country mean of `pop`, 'gdp`, `life', and 'cell', assigning to `country_stats`.  In the aggregation dictionary, start with asking for the `income` variable to get aggregated by `'first'` to get the *first* value in the partition for that country code.
2. Compute the mean of `pop`, `gdp`, and `life` for each income category by grouping based on the result from step 1.  Print the result.

```python
# Student Experimentation Cell


```

<!-- grade_id: cell-9691a58d8bf6ab53, points: 0, grade: false, solution: true -->
```python
# Solution Cell

result = io.StringIO()
with redirect_stdout(result):
    country_stats = indicators.groupby('country').agg({'income':'first','pop':'mean','gdp':'mean','life':'mean'})
    country_stats = country_stats.groupby('income').agg({'pop':'mean','gdp':'mean','life':'mean'})
    print(country_stats)
print(result.getvalue())
```

Output:
```
pop           gdp       life
income                                                   
High income           322.986667  18803.960000  78.590000
Lower middle income   709.040000   1278.145000  72.401667
Upper middle income  1378.760000  11432.326667  76.250000
```

<!-- grade_id: cell-39f8dae988e47de6, points: 2, grade: true, solution: false -->
```python
# Testing Cell

assert True

```

```python

```