# Denison CS-181/DA-210 Homework

---

## JSON Startup Exercises

```python
import json
```

**Q1** Consider the table of data:

subject | name | department
--------|------|------------------
CS      | Computer Science | MATH
DAEC    | Data Analystics (ECON) | DA
MATH    | Mathematics | MATH
ENGL    | English Literature | ENGL

This could be represented in a Python program as a dictionary mapping from a key, "subjects", to a **list of dictionaries**, one per subject. Each of these inner dictionaries map from subject, name, and department to the respective value for the row.
Write Python code to construct the above representation as a Python in-memory data structure called `ds`, and then have your code encode the data structure into a JSON-formatted string (using `dumps` appropriately), and assign to `s1` and print it. Then take `ds` and encode it as a JSON- formatted string, but include `indent=2` as an argument in the conversion, assign to `s2` and print it.

<!-- grade_id: cell-7d2d08c4ae572b41, points: 0, grade: false, solution: true -->
```python
ds = {"subjects":[{"subject":"CS", "name":"Computer Science","department":"MATH"},
                 {"subject":"DAEC", "name":"Data Analytics (ECON)","department":"DA"},
                 {"subject":"MATH", "name":"Mathematics","department":"MATH"},
                 {"subject":"ENGL", "name":"English Literature","department":"ENGL"}]}
s1 = json.dumps(ds)
s2 = json.dumps(ds, indent=2)
```

<!-- grade_id: cell-f5af9f09409c9dea, points: 2, grade: true, solution: false -->
```python
# Testing Cell
assert True
```

**Q2** Create an in-memory data structure (`ds2`) from the **same data** as in the last question.  Then add (through mutation) an additional subject to its correct location in ds2, with a subject of `"PHYS"`, name of `"Physics and Astronomy"` and department of `"PHYS"`.

Generate a file named `school0.json` **in the current directory** with the JSON-formatted data. Note that, on an open, you need to enable **write** to the created file.  Use the `indent=2` argument when encoding as JSON.

<!-- grade_id: cell-2521a3f231b53186, points: 0, grade: false, solution: true -->
```python
# Solution cell

ds2 = ds
ds2["subjects"].append({"subject":"PHYS", "name":"Physics and Astronomy","department":"PHYS"})
ds2

with open('school0.json', 'w') as outfile:
    json.dump(ds2, outfile, indent=2)
```

<!-- grade_id: cell-472cfdc98c4d15d6, points: 2, grade: true, solution: false -->
```python
# Testing cell

assert True
```

**Q3** Building on the previous question: 
    - if you successfully created `school0.json` programmatically through Q1 and Q2, use the file manager and create `school.json` as a copy of `school0.json`, or
    - create `school.json` using a text editor, and with the **same data** as the original data from Q1 with the added subject of Q2.

Write code to read and decode the JSON-formatted `school.json` file into a data structure, assigning to variable `ds3`. Once created, manipulate the `ds3` data structure by appending an additional row of data to the subjects list with subject mapping to `"BIOL"`, name mapping to `"Biology"`, and department mapping to `"BIOL"`.

<!-- grade_id: cell-b40d006ccc8ee732, points: 0, grade: false, solution: true -->
```python
# Solution cell

with open('school.json', 'w') as outfile:
    json.dump(ds2, outfile, indent=2)
    
f = open('school.json')
ds3 = json.load(f)
ds3['subjects'].append({"subject":"BIOL", "name":"Biology","department":"BIOL"})
ds3
```

Output:
```
{'subjects': [{'subject': 'CS',
   'name': 'Computer Science',
   'department': 'MATH'},
  {'subject': 'DAEC', 'name': 'Data Analytics (ECON)', 'department': 'DA'},
  {'subject': 'MATH', 'name': 'Mathematics', 'department': 'MATH'},
  {'subject': 'ENGL', 'name': 'English Literature', 'department': 'ENGL'},
  {'subject': 'PHYS', 'name': 'Physics and Astronomy', 'department': 'PHYS'},
  {'subject': 'BIOL', 'name': 'Biology', 'department': 'BIOL'}]}
```

<!-- grade_id: cell-07a2223dd9d3a417, points: 2, grade: true, solution: false -->
```python
# Testing cell

assert True
```

```python

```